Sleep and memory consolidation are linked by RNA processing genes in the Drosophila mushroom body

Author response:

The following is the authors’ response to the original reviews.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Below, I will list the points that should be addressed by the authors:

(1) Line 139: The authors conclude that the lack of a phenotype induced by knockdown of Polr1F is due to reduced baseline sleep because of the leakiness of the Genswitch system. However, it is not clear why the argument of the SybGS being leaky should not apply to all experiments done with this tool. The authors should comment on that aspect. Furthermore, this claim is testable since it should be detectable against genetic controls. An alternative explanation to the proposed scenario is that the Polr1F sleep phenotype observed in the constitutive knockdown experiment is based on developmental defects. The authors should provide additional evidence to explain the discrepancy.

We appreciate the reviewer’s insightful feedback. We assume the reviewer is referring to Regnase-1 RNAi (and not Polr1F) as Regnase-1 RNAi flies exhibit reduced sleep before dusk, potentially hindering further detection of sleep reduction. The leaky sleep reduction was based upon comparison with genetic controls in that experiment. Nevertheless, to discern whether our observations stem from developmental effects, we conducted adult-specific knockdowns of both Polr1F and Regnase-1 using the TARGET system. We generated the R35B12-Gal4:TubGal80ts line and crossed it with the UAS-Polr1FRNAi and UAS-Regnase-1RNAi lines. We confirmed that Polr1F RNAi promotes sleep when knocked down in adults (Figure 3 - supplemental figure 1). Conversely, Regnase-1 showed no effect on sleep in the adult stage, which is consistent with our nSyb-GS experiments, and suggests, as noted by the reviewer, that the Regnase-1 RNAi sleep effect is likely developmental (Figure 3 – supplemental figure 3).

(2) Line 170: Regnase1 knockdown affects all memory types, including short-term and long-term memory. The authors conclude that these genes are involved in consolidation. However, besides consolidation, it has been shown that α′β′ KCs are involved in short-term appetitive memory retrieval. Thus, an equally possible explanation is that the knockdown impairs the neuronal function per se, which would lead to a defect in all behaviors related to α′β′ KCs, rather than a specific role for consolidation. The authors have to provide additional evidence to substantiate their claim.

The exact role of Regnase-1 in the α′β′ KCs remains unclear. We acknowledge the reviewer’s concern and have amended our conclusion to include this potential explanation suggested by the reviewer.

(3) Line 87-88: For the protocol used, it was reported that GFPnls cannot be used for FACS sorting. The authors might want to comment/clarify that aspect. https://star-protocols.cell.com/protocols/1669.

For our RNA-seq experiments, we conducted single cell isolation by FACS sorting cells, instead of nuclei, labeled with GFP.nls. The protocol mentioned that GFP.nls is not effective for single nuclear RNA-seq as it is not specific for nuclei, but for our cell sorting purposes that did not matter.

(4) Line 131: The authors should report the concentration of RU486.

Sorry, this is now in methods.

(5) Line 155: Is that really 42 hours? This might be a typo. If not, it would be good to justify the prolonged re-starvation period.

Flies fed after training form sleep-dependent memories but did not show robust long-term memory after 30 h of restarvation. As starvation is a requisite for appetitive memory retrieval (Krashes and Waddell 2008), the low memory scores after 30 h could be due to inadequate starvation. Therefore, we starved flies for 42h, which is similar to the sleep-independent memory paradigm in which flies are starved for 18 h before training and then tested 24 h after training; this protocol resulted in robust long-term memory performance. These flies were fine and able to make choices in a T-maze after 42 h starvation.

(6) I will be listing mistakes/unclear points in the figures. However, all figures should be checked very carefully for clarity.

Thanks for these valuable comments. We have gone over the figures carefully and fixed any issues we found.

(7) Figure 1C: It is not entirely clear to me how this heatmap was created and what the values mean.

The 59 differentially expressed genes (DEGs) were selected based on DESeq2 described in the methods. For the heatmap, Transcripts per million (TPM) of these 59 DEGs were log-transformed and then scaled row-wise and plotted with IDEP v0.95 (http://bioinformatics.sdstate.edu/idep95/).

(8) Figures 2A and 2B: The units might be missing. For Supplementary Figure 2, it is not clear what the different groups are without looking at the main figure.

Fixed.

(9) Figure 3: The panel arrangement is confusing. Furthermore, the "B)" is cut. The same issue is present in the Supplementary Figure.

Sorry! We rearranged the panels, and fixed the issue in both figures.

(10) Figure 5B: It is not clear what the scale bar means.

Now indicated

(11) Line 119: The citation "Marygold et al n.d."?

Fixed

(12) Line 620: I'm not sure that the rate and localization of nascent peptide synthesis are measured.

Great point. We used the puromycin assay to estimate significant changes in translation. However, we did not measure the absolute translational rate or the localization of newly synthesized proteins. We rephrased this in the updated manuscript.

(13) Line 627, the authors should give the NA of the objective, further the authors should double-check the information they provide on the resolution.

Fixed, it was 20X.

(14) Line 629 "Fuji" is unclear, it might refer to the Fiji software, and in that case, it should be listed in the used software. Further, the authors have to check on the information they provide on the intensity, e.g. is that GFP fluorescence?

Yes, it was Fiji and GFP. The manuscript has been updated accordingly.

(15) Line 634, It is stated that two concentrations of CX-5461 are used, however, as far as I can see only data for the 0.2 mM.

We apologize for the confusion. Data are indeed only shown for 0.2 mM. We also tested 0.4 mM and 0.6 mM under fed conditions once and 0.1 mM under starved conditions twice. Since all effects were not significant, we only presented the complete 0.2 mM results in the supplementary figure.

(16) Line 352 "Marygold et al nd" is probably a glitch in the citation?

It’s a citation tool issue and has been fixed.

(17) The authors use apostrophe rather than a prime in describing the α "prime" β "prime" KCs

We have corrected this.

Reviewer #2 (Recommendations For The Authors):

The authors have generated an interesting study that promises to advance the understanding of how context-dependent changes in sleep and memory are executed at the molecular level. The manuscript is well-written and the statistical analyses appear robust. Major and minor comments are detailed below.

Overall, I would suggest that the authors try to obtain additional evidence that Pol1rF modulates sleep and test the effect of acute adult-stage knockdown of Polr1F and Regnase-1 specifically in ap α'β' MBNs rather than pan-neuronally.

Major comments

(1) In Figures 2 and 3 and associated supplemental figures, the authors first test for a role for Polr1F and Regnase-1 specifically in ap α'β' MBNs (Fig. 2), then test for an acute role for these proteins via pan-neuronal drug inducible expression (Fig. 3). Because the former manipulation is cell-specific and the latter is pan-neuronal, it is hard for the reader to draw conclusions pertaining to ap α'β' MBNs from the second dataset. Perhaps Regnase-1 indeed acutely regulates sleep in ap α'β' MBNs, but that effect is masked by counteracting roles in other neurons? Conversely, it remains possible that Polr1F and Regnase-1 act during development in ap α'β' MBNs to modulate sleep. Indeed, since silencing the output of ap α'β' MBNs using temperature-sensitive shibire does not alter baseline sleep (Chouhan et al., (2021) Nature), the notion that Regnase-1 could act acutely in ap α'β' MBNs to reduce baseline sleep is somewhat surprising.

The authors could address this by using a method such as TARGET (temperature-sensitive GAL80) to acutely reduce Polr1F and Regnase-1 expression specifically in ap α'β' MBNs and test how this impacts sleep.

Thanks for the very helpful suggestions. We have done the suggested experiments and discuss them above in response to Reviewer 1. They are included in the manuscript as Figure 3 – supplemental figure 1 and figure 3 – supplemental figure 3.

(2) Figure 4 presents data examining whether Polr1F and Regnase-1 knockdown suppresses training-induced increases in sleep. For the untrained flies, based on the data in Fig. 2C, E I expected that Polr1F knockdown flies would exhibit more sleep than their respective controls (Fig. 4E), but this was not the case. These data suggest that more evidence may be warranted to strengthen the link between Polr1F (and potentially Regnase-1) knockdown and sleep. Could the authors use independent RNAi constructs or cell-specific CRISPR (all available from current stock centres) to validate their current results? Related to this, it would be useful to know whether the authors outcrossed any of their transgenic reagents into a defined genetic background.

The untrained flies in figure 4E are not equivalent to flies tested for Polr1F effects on sleep in figure 2C. In Figure 4E, flies were starved for 18 h and then exposed to sucrose without an odor at ZT6. Following sucrose exposure, flies were moved to sucrose locomotor tubes, and sleep was assessed only in the ZT8-12 interval. Sleep was not significantly different between untrained R35B12>Polr1FRNAi and Polr1FRNAi/+ flies, and while it was higher in R35B12>Polr1FRNAi than in R35B12/+ untrained flies, the data overall indicate that Polr1F downregulation has no impact on sleep under these conditions and at this time. Similarly, in fully satiated settings (Figure 2C), we found no difference in sleep during the ZT8-12 period between R35B12>Polr1FRNAi flies and genetic controls. We did not outcross our transgenic lines but have now tested another available Polr1F RNAi (VDRC: v103392) (Figure 3 – supplemental figure 1). As shown in the figure, adult-specific knockdown of Polr1F by this RNAi line promoted sleep, as did the initial RNAi line.

(3) Could the authors provide additional evidence that Polr1F knockdown in ap α'β' MBNs does not enhance sleep by reducing movement? A separate assay such as climbing would be beneficial. Alternatively, examining peak activity levels at dawn/dusk from the 12L: 12D DAM data.

We checked the peak activity per minute per day for adult specific knockdown of PorlF1 and Regnase-1 (data shown in Figure 3 – supplemental figure 4). The results show that Polr1F knockdown in ap α'β' MBNs does not enhance sleep by reducing movement.

(4) In terms of validating their proposed model, over-expressing of Polr1F during appetitive training might be predicted to suppress training-induced sleep increases and potentially long-term memory. Do the authors have any evidence for this?

We were unable to find any Pol1rF overexpression line. However, we obtained the Regnase-1 over-expression line from Dr. Ryuya Fukunaga’s lab and found that Regnase-1 OE does not affect sleep (Figure 4 – supplemental figure 1).

Minor comments

(1) Abstract: can the authors please define 'ap' as anterior posterior?

Fixed.

(2) Figure 2 Supplemental 1: can the authors please denote the genotypes each color refers to in?

Fixed.

(3) In Figure 3 Supplemental 1, the authors state that acute Regnase-1 knockdown did not reduce sleep, but sleep during the night period does appear to be reduced (panel A). Was this quantified?

We quantified this, and it was not significant.

(4) Discussion, line 234: the heading of this section is 'Polr1F regulates ribosome RNA synthesis and memory' but the data presented in Figure 4 suggests that Polr1F does not affect memory. Can the authors clarify this?

We made an adjustment to the title and acknowledge that at the present time we cannot say Polr1F affects memory.

(5) Methods, Key Resource Table: can the authors please identify which fly lines were used for Polr1F and Regnase-1 knockdown experiments?

Fixed. Fly line BDSC64553 was used for Polr1F RNAi except in Figure 3 – supplemental figure 1 and 4, where VDRC 103392 was used. VDRC 27330 was used for Regnase-1 knockdown experiments.

Reviewer #3 (Recommendations For The Authors):

(1) Figure 1B: This plot is currently labelled as PCA of DEGs, which I believe is inaccurate, as such a plot is a quality control that examines the overall clustering of samples by using all read counts (not just the DEGs). In addition, the color key value of this Figure 1B is not provided.

Thank you for the insightful suggestion. The reviewer’s comment here that typically PCA plots are used for overall clustering of RNA-seq samples is indeed valid. We've acknowledged that our samples, due to their high similarity in cell populations and mild treatments, do not exhibit clear separation when we use all genes. However, we show a PathwayPCA plot of all DEGs. We aim to highlight that RNA processing pathways enriched among the DEGs account for much of the separation of the groups.

(2) A reviewer token is not provided to examine the sequencing data set.

The RNA-seq data has been submitted to the Sequence Read Archive (SRA) with NCBI BioProject accession number PRJNA1132369. The reviewer token is https://dataview.ncbi.nlm.nih.gov/object/PRJNA1132369?reviewer=cvqkddp8rjuebsjefk0f19556r.

(3) In the discussion, the author pointed out that many of the 59 DEGs have implicated functions in RNA processing. To strengthen the statement, it would be beneficial to conduct the Gene Ontology analysis to test whether the DEGs are enriched for RNA processing-related GO terms.

We have included the GO analysis results in Figure1 and another GO analysis of all DEGs in Figure 1 – supplemental figure 1.

(4) Figure 4E presents an intriguing finding because it shows that the untrained R35B12>Polr1FRNAi flies exhibit reduced sleep (instead of increased sleep) when compared to untrained Polr1/+ control flies.

Please see above response to reviewer #2 question2.

(5) For the memory assay method, the identity of odor A and odor B is not provided.

We used 4-methylcyclohexanol and 3-octanol; this information has been added into the methods section.

(6) Female flies were used for the sleep assay. However, it is not clear whether only female flies were used for the memory assay.

Mixed sexes are used for memory assays because a huge number of files is needed for these experiments. We added this information in the methods.

(7) It is important to provide olfactory acuity data on control and experimental animals to rule out that the learning/memory phenotype is caused by defects in sensing the odor used for training and testing.

Since Polr1F RNAi flies perform well, odor acuity is not an issue. Regnase1RNAi affects both short-term and long-term memories, but this seems to be a developmental issue, so we did not do the odor acuity experiments here.

(8) Line 20: "ap alpha'/beta'" neurons should be spelled as "anterior posterior (ap) alpha'/beta' neurons", as this is the first time that this anatomical name appears in this manuscript.

Fixed.

(9) Figure 2C and 2D labelling: R35B12>control; UAS control should be changed to R35B12/+ control; UAS-RNAi/+ control.

Fixed.

(10) Line 155: it is unclear why the flies were re-starved for 42hr before testing. Is this a different protocol from the 30hr re-starvation that was used by Chouhan et al., 2021?

We have explained the rationale above. The starvation period was increased to get better memory scores.

(11) Line 160: it is stated that knocking down Polr1F did not affect memory, which is consistent with Polr1f levels typically decreasing during memory consolidation. Is there a reference demonstrating that Polr1f levels typically decrease during memory consolidation?

It’s from our RNA-seq dataset from Figure1C. The level of Polr1F decreased in fed trained flies compared with other control flies.

(12) Genotype labeling in Figure 4F is inconsistent with the rest of the manuscript.

Fixed.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Memory consolidation in Drosophila can be sleep-dependent or sleep- independent, depending on the availability of food. Different regions of the mushroom body (MB) mediate these two mechanisms, with the ap α’/β’ neurons required for sleep- dependent memory consolidation in flies that are fed after training. These neurons are also involved in the increase of sleep after training, suggesting a link between sleep and memory. To better understand the mechanisms underlying sleep and memory consolidation initiation, we analyzed the transcriptome of ap α’/β’ neurons one hour after appetitive memory conditioning. A small number of genes were differentially expressed specifically in flies fed after training, but not in trained and starved flies or untrained flies. Knockdown of each of these differentially expressed genes in the ap α’/β’ neurons revealed multiple genes that affect sleep, with notable effects observed for Polr1F and Regnase-1, both of which decrease in expression after conditioning. Knockdown of Polr1F, a regulator of ribosome RNA transcription, in adult flies promotes sleep and increases pre-ribosome RNA expression as well as overall translation, supporting a function for Polr1F downregulation in memory consolidation. Conversely, knockdown of Regnase-1, an mRNA decay protein localized to the ribosome, reduces sleep. Given that Regnase-1 knockdown in ap α’/β’ neurons affects both sleep-dependent and sleep- independent memory, as well as short-term memory, Regnase-1 likely has an early role in the learning process, which may obscure a later function for its downregulation during sleep-dependent memory. These findings indicate that changes in RNA processing play a crucial role in triggering post-training sleep and memory consolidation.

Revised version: Memory consolidation in Drosophila can be sleep-dependent or sleep-independent, depending on the availability of food. Different regions of the mushroom body (MB) mediate these two mechanisms, with the anterior posterior (ap) alpha′/beta′ (α′/β′) neurons required for sleep-dependent memory consolidation in flies that are fed after training. These neurons are also involved in the increase of sleep after training, suggesting a coupling of sleep and memory. To better understand the mechanisms underlying sleep and memory consolidation initiation, we analyzed the transcriptome of ap α′/β′ neurons one hour after appetitive memory conditioning. A small number of genes were differentially expressed specifically in flies fed after training, but not in trained and starved flies or untrained flies. Knockdown of each of these differentially expressed genes in the ap α′/β′ neurons revealed multiple genes that affect sleep, with notable effects observed for Polr1F and Regnase-1, both of which decrease in expression after conditioning. Knockdown of Polr1F, a regulator of ribosome RNA transcription, in adult flies promotes sleep and increases pre-ribosome RNA expression as well as overall translation, supporting a function for Polr1F downregulation in memory consolidation. Conversely, knockdown of Regnase-1, an mRNA decay protein localized to the ribosome, reduces sleep. Given that Regnase-1 knockdown in ap α′/β′ neurons affects both sleep-dependent and sleep-independent memory, as well as short-term memory, Regnase-1 likely has an early role in the learning process, which may obscure a later function for its downregulation during sleep-dependent memory. These findings indicate that changes in RNA processing play a crucial role in triggering post-training sleep and memory consolidation.

Background

Sleep is an optimized state for memory consolidation compared to wake (Rasch and Born, 2013). Indeed, sleep is thought to reorganize and strengthen the neural connections required for novel memory formation and long-term memory consolidation, and sleep disruption leads to impaired memory consolidation (Roselli et al., 2021). Beneficial effects of sleep on memory have also been attributed to post-learning neuronal reactivation (Wagner et al., 2007; Dag et al., 2019).

Earlier version: We recently demonstrated that Drosophila switch between sleep-dependent and sleep-independent memory consolidation based on food availability (Chouhan et al., 2021). Flies that are fed after appetitive conditioning show sleep-dependent memory consolidation, which is mediated by the anterior posterior (ap) α’/β’ neurons of the mushroom body. On the other hand, flies starved after training display sleep- independent memory mediated by medial (m) neurons of the α’/β’ lobes. Neurotransmission from ap α’/β’ neurons in the first 4 hours after training in the fed flies is not only required for long term memory but also increased sleep post-training. Activation of ap α’/β’ neurons also promotes baseline sleep, supporting the idea that post training sleep is triggered by the same neurons that are required for sleep- dependent memory consolidation. However, how sleep and memory consolidation are coordinated in ap α’/β’ neurons in this time window is not understood.

Revised version: We recently demonstrated that Drosophila switch between sleep-dependent and sleep-independent memory consolidation based on food availability (Chouhan et al., 2021). Flies that are fed after appetitive conditioning show sleep-dependent memory consolidation, which is mediated by the anterior posterior (ap) α’/β’ neurons of the mushroom body. On the other hand, flies starved after training display sleep-independent memory mediated by medial (m) neurons of the α’/β’ lobes. Neurotransmission from ap α′/β′ neurons in the first 4 hours after training in the fed flies is not only required for long term memory but also increased sleep post-training. Indeed, activation of ap α′/β′ neurons promotes baseline sleep, supporting the idea that post training sleep is triggered by the same neurons that are required for sleep-dependent memory consolidation. However, how sleep and memory consolidation are coordinated in ap α′/β′ neurons in this time window is not understood.

Earlier version: Since both post-training sleep and memory consolidation occur rapidly within a few hours after training, we investigated transient gene expression changes in the ap α’/β’ neurons after training to elucidate the interplay of sleep and memory consolidation. Here, we profiled the transcriptome of ap α’/β’ neurons 1 hour after flies were fed and trained, as well as under control conditions in which flies were starved and trained or fed and untrained. We knocked down the differentially expressed genes in trained and fed flies and identified two RNA processing genes that affect sleep. Knockdown of one of these, Polr1F, a regulator of ribosomal RNA synthesis, promotes sleep and translation, both of which are required for consolidation of memory in trained and fed flies.

Revised version: Since both post-training sleep and memory consolidation occur rapidly within a few hours after training, we investigated transient gene expression changes in the ap α′/β′ neurons after training to elucidate the interplay of sleep and memory consolidation. Here, we profiled the transcriptome of ap α′/β′ neurons 1 hour after flies were fed and trained, as well as under control conditions in which flies were starved and trained or fed and untrained. We knocked down the differentially expressed genes in trained and fed flies and identified two RNA processing genes that affect sleep. Knockdown of one of these, Polr1F, a regulator of ribosomal RNA synthesis, promotes sleep and translation, both of which are required for consolidation of memory in trained and fed flies. Knockdown of Polr1F does not affect memory, which is consistent with downregulation of this gene during memory consolidation. Knockdown of Regnase-1, an mRNA decay protein, reduces sleep and disrupts short- and long-term memory, suggesting an early role, perhaps in learning. These results suggest that RNA processing is an important mechanism linking sleep to memory.

Removed in the revised version: Knockdown of Polr1F does not affect memory, which is consistent with downregulation of this gene during memory consolidation. Knockdown of Regnase-1, an mRNA decay protein, reduces sleep and disrupts short and long term memory, suggesting an early role, perhaps in learning. These results suggest that RNA processing is an important mechanism linking sleep to memory.

Results

Transient transcriptome profiling of ap α′/β′ neurons after training

Earlier version: We previously demonstrated that flies fed after appetitive memory training exhibit increased sleep and form sleep-dependent memory, which is mediated by ap α’/β’ neurons of the mushroom body (MB) α’/β’ lobes (Chouhan et al., 2021). To address the mechanisms that mediate increases in sleep and consolidation of memory in ap α’/β’ neurons, we assayed gene expression changes in ap α’/β’ neurons of flies fed after training in an appetitive conditioning paradigm. We crossed ap α’/β’ neuron driver R35B12-Gal4 (BDSC #49822) with UAS-nGFP (BDSC #4775) to label all ap α’/β’ neurons with nuclear GFP and collected 5-7 day old F1 progeny subjected to three different conditions: Trained-Fed, Trained-Starved, Untrained-Fed as illustrated (Figure 1A). The trained-starved flies served as controls for sleep-dependent changes, while flies that were fed but untrained served as controls for training-dependent changes. After one-hour, 50 mixed sex (25 for each sex) fly brains from each condition were dissected, and 500 GFP+ cells were sorted for bulk-RNA sequencing using the protocol described by Hongjie Li et al (Li et al., 2017).

Revised version: We previously demonstrated that flies fed after appetitive memory training exhibit increased sleep and form sleep-dependent memory, which is mediated by ap α′/β′ neurons of the mushroom body (MB) α’/β’ lobes (Chouhan et al., 2021). To address the mechanisms that mediate increases in sleep and consolidation of memory in ap α′/β′ neurons, we assayed gene expression changes in ap α′/β′ neurons of flies fed after training in an appetitive conditioning paradigm. We crossed ap α′/β′ neuron driver R35B12-Gal4 (BDSC #49822) with UAS-nGFP (BDSC #4775) to label all ap α′/β′ neurons with nuclear GFP and collected 5-7 day old F1 progeny, which were subjected to one of three different conditions: Trained-Fed, Trained-Starved, Untrained-Fed as illustrated (Figure 1A). The trained-starved flies served as controls for sleep-dependent changes, while flies that were fed but untrained served as controls for training-dependent changes. After one-hour, 50 mixed sex (25 for each sex) fly brains from each condition were dissected, and 500 GFP+ cells were sorted for bulk-RNA sequencing using the protocol described by Hongjie Li et al (Li et al., 2017).

Differential gene expression after training in mushroom ap α′/β′ neurons.

Earlier version: (A) 5-7 day old mixed sex wCS flies were exposed to one of the following three conditions: Trained-Fed, Trained-Starved and Untrained-Fed. Only Trained-Fed flies are expected to increase sleep after treatment and thus form sleep-dependent memory (Chouhan et al., 2021). Brain dissection, single cell suspension and cell sorting were used to extract ap α’/β’ neurons in each of these three different conditions, and bulk- sequencing of the sorted cells was conducted. (B) We sequenced four samples for each condition and principal component analysis (PCA) analysis of the differentially expressed genes (DEGs) between the three different conditions showed that the samples were separatable from each other and genes responsible for PC1, which accounted for 35% of variance, largely encode proteins involved in transcription and biosynthesis, including RNA biosynthesis, processes. (C) Heatmap of the 59 DEGs including two genes, CG13773 (Polr1F) and Regnase-1, which affect sleep and are the focus of this study. DEGs are identified by DESeq2 with the cutoff of FDR< 0.1 and fold change >1.5.

Revised version: (A) 5-7 day old mixed sex wCS flies were exposed to one of the following three conditions: Trained-Fed, Trained-Starved and Untrained-Fed. Only Trained-Fed flies are expected to increase sleep after treatment and thus form sleep-dependent memory (Chouhan et al., 2021). Brain dissection, single cell suspension and cell sorting were used to extract ap α′/β′ neurons in each of these three different conditions, and bulk-sequencing of the sorted cells was conducted. (B) We sequenced four samples for each condition and subjected them to differential gene analysis of pairwise comparison of the three conditions. We found that 59 genes are significantly different in TrainedStarved vs. TrainedFed and UntrainedFed vs. TrainedFed. (C) Gene ontology (GO) analysis of these 59 genes using FlyEnrichr (https://maayanlab.cloud/FlyEnrichr/) revealed that they encode cellular components of the 90S preribosome, Cajal body, DNA-directed RNA polymerase complex, nuclear euchromatin, and condensed chromosome.(D) Heatmap of the 59 DEGs including two genes labeled in red, CG13773 (Polr1F) and Regnase-1, which affect sleep and are the focus of this study due to their impacts on sleep. DEGs were identified by DESeq2 with the cutoff of FDR< 0.1 and fold change >1.5. Heatmaps were plotted by using TPM values of genes for each sample; data were log-transformed and scaled row-wise for visualization.

Earlier version: Our analysis of the bulk RNA-seq data revealed that most genes are not altered in ap α’/β’ neurons after one hour, so there are no significant global changes observed among the three conditions. The correlation matrix calculated using the top 75% genes showed high similarity between samples, with most values exceeding 0.9 (Figure 1 – source data 1). Nonetheless, we did observe a small subset of genes that were rapidly responsive and differentially expressed between the groups. Principal component analysis (PCA) of these differentially expressed genes (DEGs) indicated that samples from different conditions are separable, and pathway analysis of PCA indicated that transcription and RNA biosynthetic processes were influenced by the training and feeding paradigm (Figure 1B). In total, we identified 59 DEGs, of which 56 were downregulated and only three were upregulated in the Trained-Fed condition compared to the control conditions (Figure 1C, Figure 1 – source data 1). Gene ontology (GO) analysis of these 59 genes using FlyEnrichr (https://maayanlab.cloud/FlyEnrichr/) indicated that they encode cellular components of the 90S preribosome, Cajal body, DNA-directed RNA polymerase complex, nuclear euchromatin, and condensed chromosome, consistent with the PCA enrichment (Figure 1 – source data 1). Our transcriptome results suggest that ribosome biosynthesis and transcription are the initial changes in ap α’/β’ neurons of trained and fed flies.

Revised version: Our analysis of the bulk RNA-seq data revealed that most genes are not altered in ap α′/β′ neurons after one hour, so there are no significant global changes observed among the three conditions. The correlation matrix calculated using the top 75% genes showed high similarity between samples, with most values exceeding 0.9 (Figure 1 – source data). Nonetheless, we did observe a small subset of genes that were rapidly responsive and differentially expressed between the groups (Figure 1B). Pathway Principal Component Analysis (PathwayPCA) of these differentially expressed genes (DEGs), a statistical method to identify key patterns in gene expression data by reducing data dimensionality, showed that genes contributing to PC1, which accounted for 35% of the variance, largely encode proteins involved in transcription and biosynthesis, including RNA biosynthesis processes (Figure 1 – supplemental figure 1A) (Odom et al., 2020). This suggests that the training and feeding paradigm influenced transcription and RNA biosynthetic processes. Meanwhile, gene ontology analysis of these DEGs, performed using DAVID Bioinformatics (https://david.ncifcrf.gov/home.jsp), revealed significant enrichment in processes such as protein unfolding/refolding and the stress response (Figure 1 – supplemental figure 1B) (Sherman et al., 2022).

Added in the revised version: Of all the DEGs across the three groups, we focused on 59 DEGs that were significantly different in TrainedStarved vs. TrainedFed and UntrainedFed vs. TrainedFed; of these, 56 were downregulated and only three were upregulated in the Trained-Fed condition compared to the control conditions (Figure 1B, D, Figure 1 – source data). Gene ontology (GO) analysis of these 59 genes using FlyEnrichr (https://maayanlab.cloud/FlyEnrichr/) indicated that they encode cellular components of the 90S preribosome, Cajal body, DNA-directed RNA polymerase complex, nuclear euchromatin, and condensed chromosome, consistent with the PCA enrichment (Figure 1C) (Chen et al., 2013). Our transcriptome results suggest that ribosome biosynthesis and transcription are the initial changes in ap α′/β′ neurons of trained and fed flies.

Two genes expressed differentially in ap α′/β′ neurons of trained and fed flies predominately affect sleep

Earlier version: To investigate if any of the 59 DEGs identified in the ap α’/β’ neurons of trained and fed flies affect baseline sleep, we knocked down each of the genes using UAS- RNAi lines and screened for their potential effects on sleep. We used the ap α’/β’ neuron constitutive driver R35B12-Gal4 line to drive the RNAi constructs and compared the sleep patterns of knockdown flies with those of R35B12-Gal4 and UAS-RNAi control flies (Figure 2A, B). Using a cutoff of a 200 min change in sleep relative to each control group, we identified two genes, Polr1F and Regnase-1, that showed significant effects on baseline sleep. Knockdown of Polr1F, a component of RNA polymerase I complex, led to an increase in sleep. Although Polr1F has not been extensively studied in Drosophila (Marygold et al., n.d.), its human ortholog hRPA43 is part of the multi- subunit protein complex Pol I that regulates the transcription of ribosomal RNA (Beckouët et al., 2011). Knockdown of Regnase-1, an RNA-binding protein that binds to mRNA undergoing active translation and promotes mRNA decay via its ribonuclease activity, in ap α’/β’ neurons resulted in a reduction of both nighttime and daytime sleep (Figure 2A, B). Further analysis of sleep architecture revealed that knockdown of Polr1F and Regnase-1 did not significantly impact the total activity of flies, while knockdown of either Polr1F and Regnase-1 resulted in increased and decreased average length of sleep episodes, respectively (Figure 2 – supplemental figure1).

Revised version: To investigate if any of the 59 DEGs identified in the ap α′/β′ neurons of trained and fed flies affect baseline sleep, we knocked down each of the genes using UAS-RNAi lines and screened for their potential effects on sleep. We used the ap α′/β′ neuron constitutive driver R35B12-Gal4 line to drive the RNAi constructs and compared the sleep patterns of knockdown flies with those of R35B12-Gal4 and UAS-RNAi control flies (Figure 2A, B, Figure 2 – source data). Using a cutoff of a 200 min change in sleep relative to each control group, we identified two genes, Polr1F and Regnase-1, that showed significant effects on baseline sleep. Knockdown of Polr1F, a component of RNA polymerase I complex, led to an increase in sleep. Although Polr1F has not been extensively studied in Drosophila (Marygold et al., 2020), its human ortholog hRPA43 is part of the multi-subunit protein complex Pol I that regulates the transcription of ribosomal RNA (Beckouët et al., 2011). Knockdown of Regnase-1, an RNA-binding protein that binds to mRNA undergoing active translation and promotes mRNA decay via its ribonuclease activity, in ap α′/β′ neurons resulted in a reduction of both nighttime and daytime sleep (Figure 2C-F). Further analysis of sleep architecture revealed that knockdown of Polr1F and Regnase-1 did not significantly impact the total activity of flies, while knockdown of either Polr1F and Regnase-1 resulted in increased and decreased average length of sleep episodes, respectively (Figure 2 – supplemental figure1A-F).

The sleep screen of differentially expressed genes identifies Polr1F and Regnase-1 as sleep-regulating genes.

Earlier version: (A-B) Flies carrying the ap α’/β’ neuron driver R35B12-Gal4 were crossed with flies carrying UAS-RNAi constructs targeting DEGs identified from RNA-seq analysis. 5-7 days old female F1 progeny were loaded onto Trikinetics DAM monitors to measure their sleep in a 12-hour light: 12-hour dark (12:12 LD) cycle. Mean total sleep (A) and nighttime sleep (B) were calculated by Pysolo and the difference between experimental flies and Gal4 and RNAi controls was calculated separately for each independent experiment; average values comparing each experimental to its Gal4 control (X-axis) and RNAi control (Y-axis) are shown in the plots. Of all the lines screened, knockdown of Polr1F and Regnase-1 had strongest effects on sleep, producing an increase and decrease in sleep respectively. (C-F) show the representative sleep traces of R35B12-Gal4>polr1F RNAi flies and R35B12- Gal4>regnase1 RNAi. N = 23-32 per genotype from two independent replicates combined are shown in E and F respectively, and bar graphs show mean + SEM. P values for each comparison were calculated using the Kruskal-Wallis test with Dunn’s multiple comparisons test. **p<0.01, ***p<0.001, ****p<0.0001.

Revised version: (A-B) Flies carrying the ap α′/β′ neuron driver R35B12-Gal4 were crossed with flies carrying UAS-RNAi constructs targeting DEGs identified from RNA-seq analysis. 5-7 days old female F1 progeny were loaded onto Trikinetics DAM monitors to measure their sleep in a 12-hour light: 12-hour dark (12:12 LD) cycle. Mean total sleep (A) and nighttime sleep (B) were calculated by Pysolo and the difference between experimental flies and Gal4 and RNAi controls was calculated separately for each independent experiment; average values comparing each experimental to its Gal4 control (X-axis) and RNAi control (Y-axis) are shown in the plots. Of all the lines screened, knockdown of Polr1F and Regnase-1 had strongest effects on sleep, producing an increase and decrease in sleep respectively. (C-F) show the representative sleep traces of R35B12-Gal4>polr1FRNAi flies and R35B12-Gal4>regnase1RNAi. N = 23-32 per genotype from two independent replicates combined are shown in E and F respectively, and bar graphs show mean + SEM. p values for each comparison were calculated using the Kruskal-Wallis test with Dunn’s multiple comparisons test. **p<0.01, ***p<0.001, ****p<0.0001.

Earlier version: Given that these two genes reduce expression rapidly in response to training under fed conditions, we next used the inducible pan-neuronal GeneSwitch driver nSyb- GeneSwitch (GS) to determine if restricting knockdown of these two genes to the adult stage recapitulates changes in sleep and/or memory. RU486 (mifepristone) was added to normal fly food and transgene expression is induced when flies are loaded into Drosophila Activity Monitor (DAM) glass tubes (Robles-Murguia et al., 2019). We crossed the nSyb-GS flies with Polr1F or Regnase-1 RNAi flies and transferred the adult F1 progeny to RU486 tubes to induce expression of the RNAi 3-5 hours before dusk, and continuously monitored sleep for 5 days. Knockdown of Polr1F resulted in immediate inactivity and sleep (Figure 3A-C). However, with knockdown of Regnase-1, immediate changes in sleep were not noted (Figure 3 – supplemental figure 1A-C).

Revised version: Given that these two genes reduce expression rapidly in response to training under fed conditions, we next used the inducible pan-neuronal GeneSwitch driver nSyb-GeneSwitch (GS) to determine if restricting knockdown of these two genes to the adult stage recapitulates changes in sleep and/or memory. RU486 (mifepristone) was added to normal fly food and transgene expression is induced when flies are loaded into Drosophila Activity Monitor (DAM) glass tubes (Robles-Murguia et al., 2019). We crossed the nSyb-GS flies with Polr1F or Regnase-1 RNAi flies and transferred the adult F1 progeny to RU486 tubes to induce expression of the RNAi 3-5 hours before dusk, and continuously monitored sleep for 5 days. Knockdown of Polr1F resulted in immediate inactivity and sleep (Figure 3A-C). However, with knockdown of Regnase-1, immediate changes in sleep were not noted (Figure 3 – supplemental figure 3A-C). This could be due to potential leakiness of the nSyb-GS, such that it reduced sleep even without RU486 treatment, coupled with the fact that sleep before dusk is already quite low, potentially hindering detection of further sleep reduction. We also analyzed sleep for the next 3 consecutive days from day 3 to day 5 and found that constitutive pan-neuronal Polr1F knockdown increased sleep while Regnase-1 knockdown seemed to have no effect, again perhaps due to the leakiness of nSyb-GS. However, it is also possible that Regnase-1 affects sleep during developmental stages or acts differently in different brain areas (Figure 3D, E; Figure 3 – supplemental figure 3D, E). In general, these results indicate that adult specific knockdown of Polr1F promotes sleep while brain-wide adult-specific knockdown of Regnase-1 has limited effect on sleep.

Acute and chronic effects of pan-neuronal knockdown of Polr1F on sleep in adult flies.

Earlier version: (A) Representative sleep traces and transient activity plot of flies expressing Polr1F RNAi under the control of an inducible pan-neuronal driver (nSyb- GS>polr1f RNAi) with and without RU treatment. (B) Schematic representation of transient and chronic sleep measurements in nSyb-GS>polr1f RNAi flies. (C) Quantification of sleep during the first 3 hours (ZT 9-12) after F1 progeny flies were loaded into RU- or RU+ DAM tubes at ZT8-T9. Sleep was measured starting at ZT9. N = 39-40 individual flies per replicate with data from three independent replicates combined. The Mann-Whitney test was used to compare RU+ group and RU- groups. (D) Representative average sleep traces of nSyb-GS>polr1f RNAi in the RU- and RU+ DAM tubes for three consecutive days. Chronic sleep effects of pan-neuronal knockdown Polr1F were measured based on sleep data from day 3 to day 5. (E) Quantification of average total sleep of nSyb-GS>polr1f RNAi and controls in the DAM tubes from (D). Unpaired t-test was used to compare between RU- and RU+ groups. ****p<0.0001.

Revised version: (A) Schematic representation of transient and chronic sleep measurements in nSyb-GS>polr1fRNAi flies. (B) Representative sleep traces and transient activity plot of flies expressing Polr1F RNAi under the control of an inducible pan-neuronal driver (nSyb-GS>polr1f RNAi) with and without RU treatment. (C) Quantification of sleep during the first 3 hours (ZT 9-12) after F1 progeny flies were loaded into RU- or RU+ DAM tubes at ZT8-T9. Sleep was measured starting at ZT9. N = 39-40 individual flies per replicate with data from three independent replicates combined. The Mann-Whitney test was used to compare RU+ group and RU-groups. (D) Representative average sleep traces of nSyb-GS>polr1f RNAi in the RU- and RU+ DAM tubes for three consecutive days. Chronic sleep effects of pan-neuronal knockdown Polr1F were measured based on sleep data from day 3 to day 5. (E) Quantification of average total sleep of nSyb-GS>polr1fRNAi and controls in the DAM tubes from (D). Unpaired t-test was used to compare between RU- and RU+ groups. ****p<0.0001.

Added in the revised version: To further assess the adult-specific sleep function of both genes in ap α′/β′ neurons, we coupled the R35B12-Gal4 driver to tubulin-Gal80ts, which allows temperature-dependent expression of Gal4, and crossed it with the UAS-Polr1FRNAi and UAS-Regnase-1RNAi lines. We confirmed that Polr1F RNAi promotes both transient and chronic sleep when flies were transferred from permissive to restrictive temperatures during the adult stages (Figure 3 – supplemental figure 3). Conversely, Regnase-1 RNAi showed no effect on sleep in the adult stage, consistent with our nSyb-GS experiments, suggesting that the Regnase-1 RNAi sleep effect is likley developmental (Figure 3 – supplemental figure 3). We also analyzed the daily peak activity of the flies to determine if either Polr1F RNAi or Regnase-1 RNAi affected activity. We found that neither gene affected peak activity, suggesting that both genes influence sleep but not activity (Figure 3 - supplemental figure 4).

Removed in the revised version: This could be due to potential leakiness of the nSyb-GS, such that it reduced sleep even without RU486 treatment, coupled with the fact that sleep before dusk is already quite low. We also analyzed sleep for the next 3 consecutive days from day 3 to day 5 and found that constitutive pan-neuronal Polr1F knockdown increased sleep while Regnase-1 knockdown seemed to have no effect, again perhaps due to the leakiness of nSyb-GS. However, it is also possible that Regnase-1 affects sleep differently in different brain areas (Figure 3D, E; Figure 3 – supplemental figure 1 D, E). In general, these results indicate that adult specific knockdown of Polr1F promotes sleep while brain-wide adult-specific knockdown of Regnase-1 has limited effect on sleep.

Earlier version: Knockdown of Regnase-1 affects memory consolidation

Revised version: Knockdown of Regnase-1 affects memory

Earlier version: We next evaluated the impact of Polr1F and Regnase-1 knockdown on memory consolidation using our olfactory conditioning paradigm (Figure 4A). Starved flies were subjected to training to associate an odor with a reward, and then post-training, they were either kept on food vials for sleep-dependent memory consolidation or kept starved to promote sleep-independent memory consolidation. Memory tests were conducted 24 hours after training for starved flies, while fed flies were restarved for 42 hours before testing, as starvation is necessary for memory retrieval (Krashes and Waddell, 2008). We observed that constitutive knockdown of Polr1F in ap α’/β’ neurons did not affect sleep-dependent or sleep-independent memory as memory performance was comparable to that of genetic controls (Figure 4B-C). These results were consistent with the fact that Polr1F levels typically decrease during memory consolidation. Monitoring of sleep from ZT8 to ZT12 after training at zeitgeber time (ZT) 6 showed that the post-training increase in sleep was also not affected by Polr1F knockdown in ap α’/β’ neurons (Figure 4E), suggesting that acute downregulation of Polr1F may not be essential to the post-training increased sleep.

Revised version: We next evaluated the impact of Polr1F and Regnase-1 knockdown on memory consolidation using our olfactory conditioning paradigm (Figure 4A). Starved flies were subjected to training to associate an odor with a reward, and then post-training, they were either kept on food vials for sleep-dependent memory consolidation or kept starved to promote sleep-independent memory consolidation. Memory tests were conducted 24 hours after training for starved flies, while fed flies were restarved for 42 hours before testing, as starvation is necessary for memory retrieval (Krashes and Waddell, 2008). Efficacy of the RNAi line used above was verified by qPCR experiments (Figure 4 -supplemental figure 1A). We observed that constitutive knockdown of Polr1F in ap α′/β′ neurons did not affect sleep-dependent or sleep-independent memory as memory performance was comparable to that of genetic controls (Figure 4B-C). These results were consistent with the fact that Polr1F levels typically decrease during memory consolidation. Monitoring of sleep from ZT8 to ZT12 after training at zeitgeber time (ZT) 6 showed that the post-training increase in sleep was also not affected by Polr1F knockdown in ap α′/β′ neurons (Figure 4E), suggesting that acute downregulation of Polr1F may not be essential to the post-training increased sleep.

Regnase-1 expression is essential for sleep-dependent and sleep-independent memory.

Earlier version: (A) Schematic representation of the memory test protocol. (B, C) Sleep-dependent and sleep-independent memory tests were conducted under fed and starved conditions, respectively. Knockdown of Regnase-1 significantly reduces long- term memory performance in both fed and starved flies. However, knockdown of Polr1F in ap α’/β’ neurons does not affect long-term memory performance. N≥6 biological replicates, each replicate containing 100-150 flies. (D, E) Fed UAS-regnase-1-RNAi/+ and R35B12/+ flies exhibit a significant increase in sleep after training, while R35B12-Gal4>regnase-1 RNAi flies fail to show a comparable increase in post-training sleep. The total sleep in the ZT8-ZT12 interval is shown in (E). Polr1F knockdown in ap α’/β’ neurons does not affect the post-training increase in sleep. N≥32. (F) Compared to R35B12-Gal4/+ and +/UAS-Regnase-1 RNAi flies, R35B12-Gal4>reganse-1 RNAi flies show a significant decrease in the performance index in short-term memory. N≥6 biological replicates, each containing 100-150 flies. ns=not significant, p>0.05, *p<0.05, **p<0.01, ****p<0.0001.

Revised version: (A) Schematic representation of the memory test protocol. (B, C) Sleep-dependent and sleep-independent memory tests were conducted under fed and starved conditions, respectively. Knockdown of Regnase-1 significantly reduces long-term memory performance in both fed and starved flies. However, knockdown of Polr1F in ap α′/β′ neurons does not affect long-term memory performance. N≥6 biological replicates, each replicate containing 100-150 flies. (D, E) Fed UAS-regnase-1-RNAi/+ and R35B12/+ flies exhibit a significant increase in sleep after training, while R35B12-Gal4>regnase-1 RNAi flies fail to show a comparable increase in post-training sleep. The total sleep in the ZT8-ZT12 interval is shown in (E). Polr1F knockdown in ap α′/β′ neurons does not affect the post-training increase in sleep. N≥32. (F) Compared to R35B12-Gal4/+ and +/UAS-Regnase-1 RNAi flies, R35B12-Gal4>reganse-1 RNAi flies show a significant decrease in the performance index in short-term memory. N≥6 biological replicates, each containing 100-150 flies. ns=not significant, p>0.05, *p<0.05, **p<0.01, ***p<0.001, ****p<0.0001.

Earlier version: On the other hand, Regnase-1 knockdown in ap α’/β’ neurons resulted in a significant decrease in long-term memory performance in both fed and starved flies and eliminated the increase in post-training sleep (Figure 4 B-E). These findings suggested that Regnase-1 expression in ap α’/β’ neurons is necessary for both sleep-dependent and sleep-independent memory consolidation or that it is required for learning or short- term memory formation. Short-term memory tests confirmed that Regnase-1 knockdown flies performed significantly worse than control flies (Figure 4F), indicating that Regnase-1 expression is essential for the increase in post-training sleep and for short- term memory, which precedes both sleep-dependent and sleep-independent long-term memory.

Revised version: On the other hand, constitutive Regnase-1 knockdown in ap α′/β′ neurons resulted in a significant decrease in long-term memory performance in both fed and starved flies and eliminated the increase in post-training sleep (Figure 4B-E). These findings suggested that Regnase-1 expression in ap α′/β′ neurons is necessary for both sleep-dependent and sleep-independent memory consolidation. It is also possible that disruption of Regnase-1 in ap α′/β′ neurons affects learning, short-term memory formation, or appetitive memory retrieval (Aso et al., 2014; Shyu et al., 2019; Rutherford et al., 2023). Indeed, short-term memory tests confirmed that Regnase-1 knockdown flies performed significantly worse than control flies (Figure 4F), indicating that Regnase-1 acts early in the learning/memory process, which might explain its effect on sleep-dependent and sleep-independent memory. We also tested Regnase-1 overexpression for effects on sleep, and found that it does not affect sleep (Figure 4 - supplemental figure 1B) (Zhu et al., 2019).

Knockdown of Polr1F promotes translation

Since Polr1F knockdown promotes sleep and a 22 amino acid peptide within Polr1F inhibits ribosomal DNA transcription (Rothblum et al., 2014), we predicted that Polr1F acted as a suppressor for ribosomal DNA transcription, and thus knockdown of Polr1F would enhance the transcription and translation of ribosomal RNA and thereby overall protein synthesis. As Regnase-1 is thought to promote decay of mRNAs undergoing translation, its knockdown, or its downregulation following training, might also be expected to promote translation, or at least the translation of its target mRNAs. The role of protein synthesis in long-term memory consolidation is well-established across organisms (Alberini and Kandel, 2015), and nascent rRNA synthesis was also shown to be induced by training and required for memory consolidation in mouse (Allen et al., 2018). We thus used real-time qPCR of total RNA extracted from vehicle control and RU-treated fly brains to measure how precursor ribosomal RNA (Pre-rRNA) is affected by knockdown of Polr1F. We found that the pre-rRNA level increased significantly in the RU486 induction (RU+) group compared with vehicle control (RU-) group for nSyb-GS>polr1F RNAi flies (Figure 5A), indicating Polr1F knockdown results in higher levels of pre-rRNA, which is consistent with studies of Polr1F homologs in yeast cells (Thuriaux et al., 1995; Rothblum et al., 2014).

Knockdown of Polr1F results in high translation.

Earlier version: (A) The nSyb-GS>polr1f RNAi flies exhibit a significant increase in pre-rRNA levels. (B) The ex-vivo puromycin immunostaining assay was used to measure translation in dissected whole brains. The results show that knockdown of Polr1F using the pan-neuronal nSyb-GeneSwitch (GS) system increases translation relative to control flies that were not treated with RU. The normalized mean grayscales from the RU- and RU+ groups are compared using an unpaired t-test. The analysis includes data from 8-11 flies per group, with results from two independent replicates combined. (C) The schematic model illustrates the roles of Polr1F and Regnase-1 in memory consolidation. The genes Polr1F and Regnase-1 are prominently downregulated during memory consolidation in trained and fed flies, respectively. Polr1F is involved in regulating ribosomal RNA synthesis, and its decrease in levels in trained and fed flies promotes sleep and translation. In contrast, Regnase-1 is involved in mRNA decay, and its downregulation during memory consolidation may contribute to the stabilization of mRNAs that encode proteins important for long-term memory formation.

Revised version: (A) The nSyb-GS>polr1f RNAi flies exhibit a significant increase in pre-rRNA levels. (B) The ex-vivo puromycin immunostaining assay was used to measure translation in dissected whole brains. The results show that knockdown of Polr1F using the pan-neuronal nSyb-GeneSwitch (GS) system increases translation relative to control flies that were not treated with RU. The normalized mean grayscales from the RU- and RU+ groups are compared using an unpaired t-test. The analysis includes data from 8-11 flies per group, with results from two independent replicates combined. Scale bar: 100 µm. (C) The schematic model illustrates the roles of Polr1F and Regnase-1 in memory consolidation. The genes Polr1F and Regnase-1 are prominently downregulated during memory consolidation in trained and fed flies, respectively. Polr1F is involved in regulating ribosomal RNA synthesis, and its decrease in levels in trained and fed flies promotes sleep and translation. In contrast, Regnase-1 is involved in mRNA decay, and its downregulation during memory consolidation may contribute to the stabilization of mRNAs that encode proteins important for long-term memory formation.

Earlier version: The increasing ribosomal RNA should help translation and so we also used incorporation of puromycin into newly synthesized peptides as a measure of translation after inducing pan-neuronal knockdown of Polr1F. We observed high levels of translation when Porl1F was knockdown pan-neuronally (RU+) compared to the control (RU-) group (Figure 5B), indicating that Polr1F suppresses translation, probably by suppressing ribosomal RNA transcription (Rothblum et al., 2014). Thus, it may need to be downregulated after training to support translation and memory. Given that knockdown of Porl1F enhances rRNA synthesis and sleep, there is a question as to whether global alterations in rRNA synthesis impact sleep. However, feeding of rRNA inhibitor (CX5461) failed to produce rapid sleep phenotypes in both fed and starved flies (Figure 5 – supplemental figure 1), suggesting that rRNA synthesis may not directly influence sleep. Using puromycin to address effects of Regnase-1 on translation revealed an insignificant slight increase in translation (data not shown); given that Regnase-1 may specifically affect pre-existing translationally active mRNA or may act only on specific target mRNAs, its effects on de novo translation may not be obvious.

Revised version: The increasing ribosomal RNA should help translation and so we also used incorporation of puromycin into newly synthesized peptides as an estimate of translation after inducing pan-neuronal knockdown of Polr1F. We observed significantly higher levels of anti-puromycin GFP fluorescence when Porl1F was knockdown pan-neuronally (RU+) compared to the control (RU-) group (Figure 5B), indicating that Polr1F suppresses translation, probably by suppressing ribosomal RNA transcription (Rothblum et al., 2014). Thus, it may need to be downregulated after training to support translation and memory. Given that knockdown of Porl1F enhances rRNA synthesis and sleep, there is a question as to whether global alterations in rRNA synthesis impact sleep. However, feeding of an rRNA inhibitor (CX-5461) failed to produce rapid sleep phenotypes in fed or starved flies (Figure 5 – supplemental figure 5), suggesting that rRNA synthesis may not directly influence sleep. Using puromycin to address effects of Regnase-1 knockdown on translation revealed an insignificant slight increase in translation (data not shown); given that Regnase-1 may specifically affect pre-existing translationally active mRNA or may act only on specific target mRNAs, its effects on de novo translation may not be obvious.

Discussion

The anterior-posterior (ap) α′/β′ neurons of the mushroom body make critical and privileged contributions to sleep-dependent memory consolidation and post-training sleep (Krashes et al., 2007; Chouhan et al., 2021), but the mechanisms that link memory and sleep in these neurons are not known. To address this gap, we conducted transcriptomic analysis of ap α′/β′ neurons from trained and fed flies to identify genes that change rapidly under conditions that drive sleep-dependent memory. By uncovering two RNA processing genes involved in memory and sleep, our transcriptome profiling of ap α′/β′ neurons suggests that genes regulating rRNA transcription and translation are altered in the context of sleep-dependent memory consolidation.

RNA processing genes mediate sleep-dependent memory consolidation and sleep

Earlier version: Many of the 59 DEGs we identified are implicated in RNA processing. Among them, Polr1F and CG11920 affect ribosomal RNA processing, CG5654 is predicted to be part of the 90S pre-ribosome and involved in endonucleoytic cleavages (Herold et al., 2009), WDR79 encodes a small Cajal body specific RNA binding protein, Nup133 encodes a component of nuclear pore complex, Regnase-1 degrades mRNA, and CG18011, CG17568, Koko, CG11398 and Meics are all involved in RNA polymerase II- specific transcription (Di Giorgio et al., 2017; Rogg et al., 2022). The enrichment of RNA processing and translation genes is consistent with the notion that memory consolidation requires transcription and translation (Seibt and Frank, 2012; Alberini and Kandel, 2015). More importantly, our findings here indicate that alterations of some of RNA processing genes may impact sleep as well.

Revised version: Many of the 59 DEGs we identified are implicated in RNA processing. Of these, Polr1F and CG11920 affect ribosomal RNA processing, CG5654 is predicted to be part of the 90S pre-ribosome and involved in endonucleoytic cleavages (Herold et al., 2009), WDR79 encodes a small Cajal body specific RNA binding protein, Nup133 encodes a component of nuclear pore complex, Regnase-1 degrades mRNA, and CG18011, CG17568, Koko, CG11398 and Meics are all involved in RNA polymerase II-specific transcription (Di Giorgio et al., 2017; Rogg et al., 2022). The enrichment of RNA processing and translation genes is consistent with the notion that memory consolidation requires transcription and translation (Seibt and Frank, 2012; Alberini and Kandel, 2015). More importantly, our findings here indicate that alterations of some of RNA processing genes may impact sleep as well.

Earlier version: Polr1F regulates ribosome RNA synthesis and memory

Revised version: Polr1F regulates ribosome RNA synthesis and sleep

Removed in the revised version: Learning-induced changes in gene expression in memory-related neurons are often critical for long-term memory formation (Cavallaro et al., 2002; Hoedjes et al., 2015; Tadi et al., 2015). Our findings with Polr1F implicate changes in Pol I transcription during sleep-dependent memory. Polr1F(Rpa43) is predicted to be part of the RNA polymerase I complex and is involved in DNA-dependent RNA polymerase activity/rDNA transcription in yeast, especially for the initiation of ribosomal RNA (Beckouët et al., 2011; Marygold et al., 2020). As suggested by our data, Polr1F has an inhibitory role in the RNA polymerase I complex, which is also consistent with aforementioned study that found a 22 amino acid peptide within Polr1F can inhibit rDNA transcription (Rothblum et al., 2014). Based upon our finding that inducible knockdown of Polr1F rapidly promotes translation, it is likely that the rapid and dramatic decline of Polr1F after training in fed flies serves to increase de novo ribosome RNA synthesis.

Earlier version: This is consistent with the report that ribosomal RNA is induced by learning and required for memory consolidation in mice (Allen et al., 2018). While we do not know how knockdown of Polr1F promotes sleep, an attractive possibility is that higher translation is a result of elevated sleep. Sleep is thought to promote translation and it is required for sleep-dependent memory consolidation (Seibt and Frank, 2012; Chouhan et al., 2021). Alternatively, increased translation or rRNA synthesis could promote sleep. However, translation is typically thought of as a consequence of sleep rather than a cause (Zimmerman et al., 2006). Also, rRNA transcription rates remain constant throughout the day in the liver (Sinturel et al., 2017), although it is still possible that these rates vary in particular regions of the brain and affect sleep. The role of rRNA synthesis in Drosophila learning and memory has barely been explored, but our work, together with that of Allen et al. (2018), indicates that the well-known requirement for de novo protein synthesis during long-term memory consolidation (Jarome and Helmstetter, 2014) includes increased synthesis of ribosomal RNA and protein.

Revised version: Learning-induced changes in gene expression in memory-related neurons are often critical for long-term memory formation (Cavallaro et al., 2002; Hoedjes et al., 2015; Tadi et al., 2015). Our findings with Polr1F implicate changes in Pol I transcription during sleep-dependent memory. Polr1F(Rpa43) is predicted to be part of the RNA polymerase I complex and is involved in DNA-dependent RNA polymerase activity/rDNA transcription in yeast, especially for the initiation of ribosomal RNA (Beckouët et al., 2011; Marygold et al., 2020). As suggested by our data, Polr1F has an inhibitory role in the RNA polymerase I complex, which is also consistent with the aforementioned study that found a 22 amino acid peptide within Polr1F can inhibit rDNA transcription (Rothblum et al., 2014). Based upon our finding that inducible knockdown of Polr1F rapidly promotes translation, it is likely that the rapid and dramatic decline of Polr1F after training in fed flies serves to increase de novo ribosome RNA synthesis. This supports the report that ribosomal RNA is induced by learning and required for memory consolidation in mice (Allen et al., 2018). While we do not know how knockdown of Polr1F promotes sleep, an attractive possibility is that higher translation is a result of elevated sleep. Sleep is thought to promote translation and its role in memory consolidation in fed flies could be related to its effect on translation (Seibt and Frank, 2012; Chouhan et al., 2021). Alternatively, increased translation or rRNA synthesis could promote sleep. However, translation is typically thought of as a consequence of sleep rather than a cause (Zimmerman et al., 2006). Also, rRNA transcription rates remain constant throughout the day in the liver (Sinturel et al., 2017), although it is still possible that these rates vary in particular regions of the brain and affect sleep. The role of rRNA synthesis in Drosophila learning and memory has barely been explored, but our work, together with that of Allen et al. (2018), indicates that the well-known requirement for de novo protein synthesis during long-term memory consolidation (Jarome and Helmstetter, 2014) includes increased synthesis of ribosomal RNA and protein.

Earlier version: Rapid Regnase-1 inactivation may affect memory through effects on translation

Revised version: Regnase-1 inactivation may affect memory through effects on translation

Earlier version: The role of RNA binding protein Regnase-1 in the innate immune response has been extensively studied (Mino et al., 2015; Mao et al., 2017; Wei et al., 2019). However, our study sheds light on a novel function of neuronal Regase-1 in ap α’/β’ neurons on sleep and memory. Regnase-1 is an anti-inflammatory enzyme that inhibits mRNA translation during acute inflammatory responses. It localizes to the ribosomes on the surface of the endoplasmic reticulum (ER) and binds to translationally active mRNAs with specialized stem-loop structures at the 3’UTR (Uehata et al., 2013; Mino et al., 2015). When phosphorylated, Regnase-1 is released from the ER (Tanaka et al., 2019). Because functional Regnase-1 binds and degrades its bound mRNA, Regnase-1 inactivation leads to an increase of its target mRNA (Uehata et al., 2013). The target mRNAs of Regnase-1 in immune cells encode proinflammatory cytokines, which can then be expressed when Regnase-1 is inactivated. However, Regnase-1 has also been reported to modulate cytokines and neuronal injury in the microglia in rats (Liu et al., 2016). Regulation of Regnase-1 is usually rapid and transient, and its rapid response to microenvironmental changes, different pathological states and stress are critical for cellular adaption (Mao et al., 2017).

Revised version: The role of RNA binding protein Regnase-1 in the innate immune response has been extensively studied (Mino et al., 2015; Mao et al., 2017; Wei et al., 2019). However, our study sheds light on a novel function of Regase-1 in ap α′/β′ neurons on sleep and memory. Regnase-1 is an anti-inflammatory enzyme that inhibits mRNA translation during acute inflammatory responses. It localizes to the ribosomes on the surface of the endoplasmic reticulum (ER) and binds to translationally active mRNAs with specialized stem-loop structures at the 3′UTR (Uehata et al., 2013; Mino et al., 2015). When phosphorylated, Regnase-1 is released from the ER (Tanaka et al., 2019). Because functional Regnase-1 binds and degrades its bound mRNA, Regnase-1 inactivation leads to an increase of its target mRNA (Uehata et al., 2013). The target mRNAs of Regnase-1 in immune cells encode proinflammatory cytokines, which can then be expressed when Regnase-1 is inactivated. However, Regnase-1 has also been reported to modulate cytokines and neuronal injury in the microglia in rats (Liu et al., 2016). Regulation of Regnase-1 is usually rapid and transient, and its rapid response to microenvironmental changes, different pathological states and stress is critical for cellular adaption (Mao et al., 2017).

Earlier version: Our study reveals that the expression of Regnase-1 changes after training, and constitutive downregulation of Regnase-1 in ap α’/β’ neurons reduces sleep and causes deficits in learning and memory consolidation. We suggest that downregulation of Regnase-1 following training in fed flies may be crucial for the promotion of long-term memory, perhaps by promoting translation of specific transcripts. For instance, downregulation of Regnase-1 after training may release a pool of mRNAs that are normally targeted by it for decay and that are important for consolidation of memory. On the other hand, we find that constitutive loss of Regnase-1 impairs sleep-independent memory and also short-term memory, suggesting that it is required early in the learning and memory process, likely for learning. Interestingly, constitutive knockdown of Regnase-1 also reduces sleep and prevents sleep increase after training. The requirement for sleep is linked most to long-term memory rather than to learning, but increases in neural activity after training and cognitive tasks are thought to increase the need to sleep (Wagner et al., 2007; Diekelmann and Born, 2010), so it is possible that the reduced sleep phenotype of the constitutive Regnase-1 knockdown is related to its role in learning.

Revised version: Our study reveals that the expression of Regnase-1 changes after training, and constitutive downregulation of Regnase-1 in ap α′/β′ neurons reduce sleep and causes deficits in learning and memory consolidation. We suggest that downregulation of Regnase-1 following training in fed flies may be crucial for the promotion of long-term memory, perhaps by promoting translation of specific transcripts. For instance, downregulation of Regnase-1 after training may release a pool of mRNAs that are normally targeted by it for decay and that are important for consolidation of memory. On the other hand, we find that constitutive loss of Regnase-1 impairs sleep-independent memory and also short-term memory, suggesting that it is required early in the learning and memory process, likely for learning. Interestingly, constitutive knockdown of Regnase-1 also reduces sleep and prevents sleep increase after training. However, our adult specific knockdown of Regnase-1 does not affect sleep, suggesting that expression of Regnase-1 during development is important for both memory and sleep. Given that ap α′/β′ neurons have a role in several aspects of memory, consolidation as well as short-term appetitive memory retrieval, how Regnase-1 acts within these neurons to regulate memory clearly requires more investigation.

Earlier version: Our study of local molecular changes in ap α’/β’ neurons after training highlights how ribosomal RNA transcription and mRNA translation work in concert during the consolidation of sleep-dependent memory. Neurons are regulated to generate short bursts of boosted transcription and boosted translation, which are required for new protein synthesis. How sleep is involved in this boosted protein synthesis process is unclear, but we suggest that sleep affects several steps of the central dogma, and that deleterious consequences of sleep deprivation are also mediated through effects on rRNA transcription and translation. Overall, our findings indicate the importance of maintaining RNA homeostasis for sleep and memory.

Revised version: Our study of local molecular changes in ap α′/β′ neurons after training highlights how ribosomal RNA transcription and mRNA translation work in concert during the consolidation of sleep-dependent memory. Neurons are regulated to generate short bursts of boosted transcription and boosted translation, which are required for new protein synthesis. How sleep is involved in this boosted protein synthesis process is unclear, but we suggest that sleep affects several steps of the central dogma, and that deleterious consequences of sleep deprivation are mediated, at least in part, through effects on rRNA transcription and translation. Overall, our findings indicate the importance of maintaining RNA homeostasis for sleep and memory.

Added in the revised version: Supporting information